Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Nguyễn Duy Đạt"
COLLABORATORS = "CS431.N21"
```

---

# LUYỆN TẬP OOP - PHẦN 1

Viết chương trình cài đặt 3 lớp (class) sau: `Geometry` , `Square`, và `Round`. Cho biết 2 class `Square` và `Round` kế thừa từ class `Geometry`.

Cả 3 class đều có những thuộc tính, phương thức sau:

- Thuộc tính:
    - `property` (`float`): có thể xem là độ dài cạnh (đối với hình vuông) hoặc bán kính (đối với hình tròn).

- Phương thức:
    - Hàm khởi tạo `__init__`.
    - Hàm tính diện tích `calculate_area`.

Ví dụ về hàm `calculate_area` của class `Square`:
- Thuộc tính của object:
```
property = 2
```
- Input:
```
# no input
```
- Output:
```
4
```

Ví dụ về hàm `calculate_area` của class `Round`:
- Thuộc tính của object:
```
property = 2
```
- Input:
```
# no input
```
- Output:
```
12.56637
```

**Lưu ý:** Dùng `math.pi` để tính diện tích, kết quả làm tròn 5 chữ số thập phân.

**Yêu cầu 1:** Cài đặt class `Geometry`. Cài đặt phương thức `calculate_area` là abstract method.


<!-- grade_id: cell-a2a977de6573d46c, points: 0, grade: false, solution: true -->
```python
from abc import abstractmethod


# YOUR CODE HERE
class Geometry:
    def __init__(self, property: float):
        self.property = property
    
    @abstractmethod
    def calculate_area(self):
        return self.property
        
    # raise NotImplementedError()
```

<!-- grade_id: cell-cd8894a81325a585, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 1
geometry_obj = Geometry(2)
assert geometry_obj.property == 2
```

<!-- grade_id: cell-e5f344956ad1440a, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 2
```

**Yêu cầu 2:** Cài đặt class `Square` kế thừa từ class `Geometry`.


<!-- grade_id: cell-735e1d5e4bdbe1e7, points: 0, grade: false, solution: true -->
```python
class Square(Geometry):
    def __init__(self, property):
        # gọi hàm __init__ của lớp cha bằng hàm `super()`
        # YOUR CODE HERE
        super(Square, self).__init__(property)
        # raise NotImplementedError()

    # cài đặt hàm `calculate_area`
    # YOUR CODE HERE
    def calculate_area(self):
        return self.property ** 2
    # raise NotImplementedError()
```

<!-- grade_id: cell-7d81c0737ff8ff5c, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 3
square_obj = Square(2)
assert square_obj.calculate_area() == 4
```

<!-- grade_id: cell-7483770bd88dec92, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 4
```

**Yêu cầu 3:** Cài đặt class `Round` kế thừa từ class `Geometry`. Lưu ý sử dụng `math.pi` và làm tròn kết quả 5 chữ số thập phân.

<!-- grade_id: cell-43e8069b1597af6b, points: 0, grade: false, solution: true -->
```python
import math

class Round(Geometry):
    def __init__(self, property):
        # gọi hàm __init__ của lớp cha bằng hàm `super()`
        # YOUR CODE HERE
        super(Round, self).__init__(property)
        # raise NotImplementedError()

    # cài đặt hàm `calculate_area`
    # YOUR CODE HERE
    def calculate_area(self):
        return round(math.pi * (self.property ** 2), 5)
    # raise NotImplementedError()
```

<!-- grade_id: cell-3e614da348102a27, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 5
square_obj = Round(2)
assert square_obj.calculate_area() == 12.56637
```

<!-- grade_id: cell-86742bf27acceec1, points: 5, grade: true, solution: false -->
```python
# cell chấm điểm 6
```